# Week 3 Problem Set

## Homeworks

**HW1.** *Fibonaci:* Write a function to find the Fibonacci number given an index. These are the example of a the first few numbers in Fibonacci series.

| Indices    | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7  | 8  | 9  |
|------------|---|---|---|---|---|---|---|----|----|----|
| The series | 0 | 1 | 1 | 2 | 3 | 5 | 8 | 13 | 21 | 34 |

We can write that the Fibonacci number at index *i* is given by:

$fibo(i) = fibo(i-1) + fibo(i-2)$

User recursion for your implementation.

```python
def fibonacci(index):
    if index == 0:
        return 0
    elif index == 1:
        return 1
    else:
        return fibonacci(index - 1) + fibonacci(index - 2)
    pass
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(3) == 2
assert fibonacci(7) == 13
assert fibonacci(9) == 34
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Max-Heapify:* Recall the algorithm for max-heapify in restoring the *max-heap property* of a binary heap tree. In the previous implementation, we used iteration. Implement the same algorithm using recursion. 

```python
def left_of(index):
    return ((index * 2) + 1)
    pass

def right_of(index):
    return ((index + 1) * 2)
    pass

def max_heapify(array, index, heap_size):
    if left_of(index) >= heap_size:
        return None
    elif right_of(index) >= heap_size:
        if array[left_of(index)] > array[index]:
            array[index], array[left_of(index)] = array[left_of(index)], array[index]
    else:
        if array[left_of(index)] > array[right_of(index)] and array[left_of(index)] > array[index]:
            array[index], array[left_of(index)] = array[left_of(index)], array[index]
            max_heapify(array, left_of(index), heap_size)
        elif array[right_of(index)] > array[index]:
            array[index], array[right_of(index)] = array[right_of(index)], array[index]
            max_heapify(array, right_of(index), heap_size)
    pass
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *String Permutation:* Write two functions to return an array of all the permutations of a string. For example, if the input is `abc`, the output should be

```
output = ["abc", "acb", "bac", "bca", "cab", "cba"]
```

The first function only has one argument which is the input string, i.e. `permutate(s)`. The second function is the recursive function with two arguments String 1 (`str1`) and and String 2 (`str2`). The first function calls the second method `permutate("", s)` at the beginning. The second function should use a loop to move a character from `str2` to `str1` and recursively invokes it with a new `str1` and `str2`. 

```python
def permutate(s):
    return permutate_recursive("", s)
    pass

def permutate_recursive(str1, str2):
    result = []
    
    if len(str2) == 1:
        return [str1 + str2]
    
    for i in range(-1,-(len(str2) + 1), -1):        
        # Start with str1 = "c", therefore to put cba at the back, result has to be at the back
        result = permutate_recursive(str1 + str2[i], str2.replace(str2[i], "")) + result
        
    return result
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert permutate("abc")==["abc", "acb", "bac", "bca", "cab", "cba"]
assert permutate("abcd")==["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"]
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *GCD:* Implement the recursive implementation of Euclid's algorithm for Greatest Common Divisor (GCD).

Hint: [Wikipedia: Greatest Common Divisor](https://en.wikipedia.org/wiki/Greatest_common_divisor#Euclid's_algorithm)

```python
def gcd(a, b):
    if a == b or b == 0:
        return a
    elif a == 0:
        return b
    elif a > b:
        return gcd(a - b, b)
    else:
        return gcd(a, b - a)
    pass
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert gcd(0,0) == 0
assert gcd(2,0) == 2
assert gcd(0,2) == 2
assert gcd(2,3) == 1
assert gcd(4, 28) == 4
assert gcd(1024, 526) == 2
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```